<?xml version="1.0" encoding="UTF-8"?>
<quiz><question type="coderunner"><name><text>Q1 — Nombre de nœuds</text></name><questiontext format="html"><text>&lt;div style="background:#e4f1f2;border:1px solid #b7dadd;border-left:5px solid #026573;border-radius:8px;padding:20px 24px;margin-bottom:18px;"&gt;&lt;p style="color:#026573;font-family:Georgia,serif;font-size:13px;font-weight:bold;letter-spacing:2px;text-transform:uppercase;margin:0 0 10px 0;"&gt;Q1 &amp;mdash; Nombre de nœuds&lt;/p&gt;&lt;p style="color:#1e293b;font-family:Georgia,serif;font-size:15px;line-height:1.8;margin:0;"&gt;On représente un arbre binaire par des listes imbriquées.&lt;/p&gt;&lt;/div&gt;&lt;p style="font-family:Georgia,serif;font-size:15px;color:#1e293b;line-height:1.8;"&gt;Écrire une fonction &lt;strong&gt;&lt;code&gt;nb_noeuds(a)&lt;/code&gt;&lt;/strong&gt; qui renvoie le nombre de nœuds de l'arbre &lt;code&gt;a&lt;/code&gt;.&lt;/p&gt;&lt;div style="background:#f4f9fc;border:1px solid #b7dadd;border-left:5px solid #026573;border-radius:8px;padding:14px 20px;margin:14px 0;"&gt;&lt;pre style="color:#0a3a40;font-family:'Courier New',monospace;font-size:14px;margin:0;background:none;border:none;"&gt;nb_noeuds([4,[2,None,None],[6,None,None]]) &amp;rarr; 3&lt;/pre&gt;&lt;/div&gt;&lt;p style="font-family:Georgia,serif;font-size:14px;color:#475569;line-height:1.7;"&gt;Convention : un arbre vaut &lt;code&gt;None&lt;/code&gt; (vide) ou &lt;code&gt;[valeur, gauche, droite]&lt;/code&gt;. Ex. &lt;code&gt;[4, [2, None, None], [6, None, None]]&lt;/code&gt;.&lt;/p&gt;</text></questiontext><generalfeedback format="html"><text /></generalfeedback><defaultgrade>2.0000000</defaultgrade><penalty>0.0000000</penalty><hidden>0</hidden><idnumber /><coderunnertype>python3</coderunnertype><prototypetype>0</prototypetype><allornothing>0</allornothing><penaltyregime>0</penaltyregime><precheck>0</precheck><hidecheck>0</hidecheck><showsource>0</showsource><answerboxlines>18</answerboxlines><answerboxcolumns>100</answerboxcolumns><answerpreload>def nb_noeuds(a):
    # Compléter : cas de base puis appel récursif
    ...
</answerpreload><globalextra /><useace /><resultcolumns /><template /><iscombinatortemplate /><allowmultiplestdins /><answer>def nb_noeuds(a):
    if a is None:
        return 0
    return 1 + nb_noeuds(a[1]) + nb_noeuds(a[2])
</answer><validateonsave>1</validateonsave><testsplitterre /><language /><acelang /><sandbox /><grader /><cputimelimitsecs /><memlimitmb /><sandboxparams /><templateparams /><hoisttemplateparams>1</hoisttemplateparams><extractcodefromjson>1</extractcodefromjson><templateparamslang>None</templateparamslang><templateparamsevalpertry>0</templateparamsevalpertry><templateparamsevald>{}</templateparamsevald><twigall>0</twigall><uiplugin /><uiparameters /><attachments>0</attachments><attachmentsrequired>0</attachmentsrequired><maxfilesize>10240</maxfilesize><filenamesregex /><filenamesexplain /><displayfeedback>1</displayfeedback><giveupallowed>0</giveupallowed><prototypeextra /><testcases><testcase testtype="0" useasexample="1" hiderestiffail="0" mark="1.0000000"><testcode><text>a = [4, [2, [1,None,None], [3,None,None]], [6, [5,None,None], [7,None,None]]]
print(nb_noeuds(a))
</text></testcode><stdin><text /></stdin><expected><text>7
</text></expected><extra><text /></extra><display><text>SHOW</text></display></testcase><testcase testtype="0" useasexample="1" hiderestiffail="0" mark="1.0000000"><testcode><text>print(nb_noeuds(None))
</text></testcode><stdin><text /></stdin><expected><text>0
</text></expected><extra><text /></extra><display><text>SHOW</text></display></testcase><testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000"><testcode><text>print(nb_noeuds([5,None,None]))
</text></testcode><stdin><text /></stdin><expected><text>1
</text></expected><extra><text /></extra><display><text>SHOW</text></display></testcase><testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000"><testcode><text>print(nb_noeuds([1,[2,None,None],None]))
</text></testcode><stdin><text /></stdin><expected><text>2
</text></expected><extra><text /></extra><display><text>HIDE</text></display></testcase></testcases></question><question type="coderunner"><name><text>Q2 — Hauteur</text></name><questiontext format="html"><text>&lt;div style="background:#e4f1f2;border:1px solid #b7dadd;border-left:5px solid #026573;border-radius:8px;padding:20px 24px;margin-bottom:18px;"&gt;&lt;p style="color:#026573;font-family:Georgia,serif;font-size:13px;font-weight:bold;letter-spacing:2px;text-transform:uppercase;margin:0 0 10px 0;"&gt;Q2 &amp;mdash; Hauteur de l'arbre&lt;/p&gt;&lt;p style="color:#1e293b;font-family:Georgia,serif;font-size:15px;line-height:1.8;margin:0;"&gt;La hauteur est la longueur du plus long chemin de la racine à une feuille.&lt;/p&gt;&lt;/div&gt;&lt;p style="font-family:Georgia,serif;font-size:15px;color:#1e293b;line-height:1.8;"&gt;Écrire &lt;strong&gt;&lt;code&gt;hauteur(a)&lt;/code&gt;&lt;/strong&gt; renvoyant la hauteur de &lt;code&gt;a&lt;/code&gt;, avec la convention : arbre vide &amp;rarr; &lt;code&gt;-1&lt;/code&gt;, feuille seule &amp;rarr; &lt;code&gt;0&lt;/code&gt;.&lt;/p&gt;&lt;div style="background:#f4f9fc;border:1px solid #b7dadd;border-left:5px solid #026573;border-radius:8px;padding:14px 20px;margin:14px 0;"&gt;&lt;pre style="color:#0a3a40;font-family:'Courier New',monospace;font-size:14px;margin:0;background:none;border:none;"&gt;hauteur([4,[2,None,None],[6,None,None]]) &amp;rarr; 1&lt;/pre&gt;&lt;/div&gt;&lt;p style="font-family:Georgia,serif;font-size:14px;color:#475569;line-height:1.7;"&gt;Convention : un arbre vaut &lt;code&gt;None&lt;/code&gt; (vide) ou &lt;code&gt;[valeur, gauche, droite]&lt;/code&gt;. Ex. &lt;code&gt;[4, [2, None, None], [6, None, None]]&lt;/code&gt;.&lt;/p&gt;</text></questiontext><generalfeedback format="html"><text /></generalfeedback><defaultgrade>2.0000000</defaultgrade><penalty>0.0000000</penalty><hidden>0</hidden><idnumber /><coderunnertype>python3</coderunnertype><prototypetype>0</prototypetype><allornothing>0</allornothing><penaltyregime>0</penaltyregime><precheck>0</precheck><hidecheck>0</hidecheck><showsource>0</showsource><answerboxlines>18</answerboxlines><answerboxcolumns>100</answerboxcolumns><answerpreload>def hauteur(a):
    # Convention : arbre vide -&gt; -1
    ...
</answerpreload><globalextra /><useace /><resultcolumns /><template /><iscombinatortemplate /><allowmultiplestdins /><answer>def hauteur(a):
    if a is None:
        return -1
    return 1 + max(hauteur(a[1]), hauteur(a[2]))
</answer><validateonsave>1</validateonsave><testsplitterre /><language /><acelang /><sandbox /><grader /><cputimelimitsecs /><memlimitmb /><sandboxparams /><templateparams /><hoisttemplateparams>1</hoisttemplateparams><extractcodefromjson>1</extractcodefromjson><templateparamslang>None</templateparamslang><templateparamsevalpertry>0</templateparamsevalpertry><templateparamsevald>{}</templateparamsevald><twigall>0</twigall><uiplugin /><uiparameters /><attachments>0</attachments><attachmentsrequired>0</attachmentsrequired><maxfilesize>10240</maxfilesize><filenamesregex /><filenamesexplain /><displayfeedback>1</displayfeedback><giveupallowed>0</giveupallowed><prototypeextra /><testcases><testcase testtype="0" useasexample="1" hiderestiffail="0" mark="1.0000000"><testcode><text>a = [4, [2, [1,None,None], [3,None,None]], [6, [5,None,None], [7,None,None]]]
print(hauteur(a))
</text></testcode><stdin><text /></stdin><expected><text>2
</text></expected><extra><text /></extra><display><text>SHOW</text></display></testcase><testcase testtype="0" useasexample="1" hiderestiffail="0" mark="1.0000000"><testcode><text>print(hauteur([5,None,None]))
</text></testcode><stdin><text /></stdin><expected><text>0
</text></expected><extra><text /></extra><display><text>SHOW</text></display></testcase><testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000"><testcode><text>print(hauteur(None))
</text></testcode><stdin><text /></stdin><expected><text>-1
</text></expected><extra><text /></extra><display><text>SHOW</text></display></testcase><testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000"><testcode><text>print(hauteur([1,[2,[3,None,None],None],None]))
</text></testcode><stdin><text /></stdin><expected><text>2
</text></expected><extra><text /></extra><display><text>HIDE</text></display></testcase></testcases></question><question type="coderunner"><name><text>Q3 — Parcours infixe</text></name><questiontext format="html"><text>&lt;div style="background:#e4f1f2;border:1px solid #b7dadd;border-left:5px solid #026573;border-radius:8px;padding:20px 24px;margin-bottom:18px;"&gt;&lt;p style="color:#026573;font-family:Georgia,serif;font-size:13px;font-weight:bold;letter-spacing:2px;text-transform:uppercase;margin:0 0 10px 0;"&gt;Q3 &amp;mdash; Parcours infixe&lt;/p&gt;&lt;p style="color:#1e293b;font-family:Georgia,serif;font-size:15px;line-height:1.8;margin:0;"&gt;Le parcours infixe visite : sous-arbre gauche, racine, sous-arbre droit.&lt;/p&gt;&lt;/div&gt;&lt;p style="font-family:Georgia,serif;font-size:15px;color:#1e293b;line-height:1.8;"&gt;Écrire &lt;strong&gt;&lt;code&gt;infixe(a)&lt;/code&gt;&lt;/strong&gt; qui renvoie la &lt;strong&gt;liste&lt;/strong&gt; des valeurs dans l'ordre infixe.&lt;/p&gt;&lt;div style="background:#f4f9fc;border:1px solid #b7dadd;border-left:5px solid #026573;border-radius:8px;padding:14px 20px;margin:14px 0;"&gt;&lt;pre style="color:#0a3a40;font-family:'Courier New',monospace;font-size:14px;margin:0;background:none;border:none;"&gt;infixe([4,[2,None,None],[6,None,None]]) &amp;rarr; [2, 4, 6]&lt;/pre&gt;&lt;/div&gt;&lt;p style="font-family:Georgia,serif;font-size:14px;color:#475569;line-height:1.7;"&gt;Convention : un arbre vaut &lt;code&gt;None&lt;/code&gt; (vide) ou &lt;code&gt;[valeur, gauche, droite]&lt;/code&gt;. Ex. &lt;code&gt;[4, [2, None, None], [6, None, None]]&lt;/code&gt;.&lt;/p&gt;</text></questiontext><generalfeedback format="html"><text /></generalfeedback><defaultgrade>2.0000000</defaultgrade><penalty>0.0000000</penalty><hidden>0</hidden><idnumber /><coderunnertype>python3</coderunnertype><prototypetype>0</prototypetype><allornothing>0</allornothing><penaltyregime>0</penaltyregime><precheck>0</precheck><hidecheck>0</hidecheck><showsource>0</showsource><answerboxlines>18</answerboxlines><answerboxcolumns>100</answerboxcolumns><answerpreload>def infixe(a):
    # gauche, puis racine, puis droite
    ...
</answerpreload><globalextra /><useace /><resultcolumns /><template /><iscombinatortemplate /><allowmultiplestdins /><answer>def infixe(a):
    if a is None:
        return []
    return infixe(a[1]) + [a[0]] + infixe(a[2])
</answer><validateonsave>1</validateonsave><testsplitterre /><language /><acelang /><sandbox /><grader /><cputimelimitsecs /><memlimitmb /><sandboxparams /><templateparams /><hoisttemplateparams>1</hoisttemplateparams><extractcodefromjson>1</extractcodefromjson><templateparamslang>None</templateparamslang><templateparamsevalpertry>0</templateparamsevalpertry><templateparamsevald>{}</templateparamsevald><twigall>0</twigall><uiplugin /><uiparameters /><attachments>0</attachments><attachmentsrequired>0</attachmentsrequired><maxfilesize>10240</maxfilesize><filenamesregex /><filenamesexplain /><displayfeedback>1</displayfeedback><giveupallowed>0</giveupallowed><prototypeextra /><testcases><testcase testtype="0" useasexample="1" hiderestiffail="0" mark="1.0000000"><testcode><text>a = [4, [2, [1,None,None], [3,None,None]], [6, [5,None,None], [7,None,None]]]
print(infixe(a))
</text></testcode><stdin><text /></stdin><expected><text>[1, 2, 3, 4, 5, 6, 7]
</text></expected><extra><text /></extra><display><text>SHOW</text></display></testcase><testcase testtype="0" useasexample="1" hiderestiffail="0" mark="1.0000000"><testcode><text>print(infixe([5,None,None]))
</text></testcode><stdin><text /></stdin><expected><text>[5]
</text></expected><extra><text /></extra><display><text>SHOW</text></display></testcase><testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000"><testcode><text>print(infixe(None))
</text></testcode><stdin><text /></stdin><expected><text>[]
</text></expected><extra><text /></extra><display><text>SHOW</text></display></testcase><testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000"><testcode><text>print(infixe([1,[2,None,None],[3,None,None]]))
</text></testcode><stdin><text /></stdin><expected><text>[2, 1, 3]
</text></expected><extra><text /></extra><display><text>HIDE</text></display></testcase></testcases></question><question type="coderunner"><name><text>Q4 — Recherche dans un ABR</text></name><questiontext format="html"><text>&lt;div style="background:#e4f1f2;border:1px solid #b7dadd;border-left:5px solid #026573;border-radius:8px;padding:20px 24px;margin-bottom:18px;"&gt;&lt;p style="color:#026573;font-family:Georgia,serif;font-size:13px;font-weight:bold;letter-spacing:2px;text-transform:uppercase;margin:0 0 10px 0;"&gt;Q4 &amp;mdash; Recherche dans un ABR&lt;/p&gt;&lt;p style="color:#1e293b;font-family:Georgia,serif;font-size:15px;line-height:1.8;margin:0;"&gt;Dans un arbre binaire de recherche (ABR) : gauche &amp;lt; nœud &amp;lt; droite.&lt;/p&gt;&lt;/div&gt;&lt;p style="font-family:Georgia,serif;font-size:15px;color:#1e293b;line-height:1.8;"&gt;Écrire &lt;strong&gt;&lt;code&gt;recherche(a, x)&lt;/code&gt;&lt;/strong&gt; renvoyant &lt;code&gt;True&lt;/code&gt; si &lt;code&gt;x&lt;/code&gt; figure dans l'ABR &lt;code&gt;a&lt;/code&gt;, &lt;code&gt;False&lt;/code&gt; sinon. Exploiter la propriété d'ordre pour ne descendre que d'un côté.&lt;/p&gt;&lt;div style="background:#f4f9fc;border:1px solid #b7dadd;border-left:5px solid #026573;border-radius:8px;padding:14px 20px;margin:14px 0;"&gt;&lt;pre style="color:#0a3a40;font-family:'Courier New',monospace;font-size:14px;margin:0;background:none;border:none;"&gt;recherche([4,[2,None,None],[6,None,None]], 6) &amp;rarr; True&lt;/pre&gt;&lt;/div&gt;&lt;p style="font-family:Georgia,serif;font-size:14px;color:#475569;line-height:1.7;"&gt;Convention : un arbre vaut &lt;code&gt;None&lt;/code&gt; (vide) ou &lt;code&gt;[valeur, gauche, droite]&lt;/code&gt;. Ex. &lt;code&gt;[4, [2, None, None], [6, None, None]]&lt;/code&gt;.&lt;/p&gt;</text></questiontext><generalfeedback format="html"><text /></generalfeedback><defaultgrade>2.0000000</defaultgrade><penalty>0.0000000</penalty><hidden>0</hidden><idnumber /><coderunnertype>python3</coderunnertype><prototypetype>0</prototypetype><allornothing>0</allornothing><penaltyregime>0</penaltyregime><precheck>0</precheck><hidecheck>0</hidecheck><showsource>0</showsource><answerboxlines>18</answerboxlines><answerboxcolumns>100</answerboxcolumns><answerpreload>def recherche(a, x):
    # comparer x a la racine puis descendre du bon cote
    ...
</answerpreload><globalextra /><useace /><resultcolumns /><template /><iscombinatortemplate /><allowmultiplestdins /><answer>def recherche(a, x):
    if a is None:
        return False
    if x == a[0]:
        return True
    if x &lt; a[0]:
        return recherche(a[1], x)
    return recherche(a[2], x)
</answer><validateonsave>1</validateonsave><testsplitterre /><language /><acelang /><sandbox /><grader /><cputimelimitsecs /><memlimitmb /><sandboxparams /><templateparams /><hoisttemplateparams>1</hoisttemplateparams><extractcodefromjson>1</extractcodefromjson><templateparamslang>None</templateparamslang><templateparamsevalpertry>0</templateparamsevalpertry><templateparamsevald>{}</templateparamsevald><twigall>0</twigall><uiplugin /><uiparameters /><attachments>0</attachments><attachmentsrequired>0</attachmentsrequired><maxfilesize>10240</maxfilesize><filenamesregex /><filenamesexplain /><displayfeedback>1</displayfeedback><giveupallowed>0</giveupallowed><prototypeextra /><testcases><testcase testtype="0" useasexample="1" hiderestiffail="0" mark="1.0000000"><testcode><text>a = [4, [2, [1,None,None], [3,None,None]], [6, [5,None,None], [7,None,None]]]
print(recherche(a, 5))
</text></testcode><stdin><text /></stdin><expected><text>True
</text></expected><extra><text /></extra><display><text>SHOW</text></display></testcase><testcase testtype="0" useasexample="1" hiderestiffail="0" mark="1.0000000"><testcode><text>a = [4, [2, [1,None,None], [3,None,None]], [6, [5,None,None], [7,None,None]]]
print(recherche(a, 8))
</text></testcode><stdin><text /></stdin><expected><text>False
</text></expected><extra><text /></extra><display><text>SHOW</text></display></testcase><testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000"><testcode><text>print(recherche(None, 1))
</text></testcode><stdin><text /></stdin><expected><text>False
</text></expected><extra><text /></extra><display><text>SHOW</text></display></testcase><testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000"><testcode><text>a = [4, [2, [1,None,None], [3,None,None]], [6, [5,None,None], [7,None,None]]]
print(recherche(a, 1))
</text></testcode><stdin><text /></stdin><expected><text>True
</text></expected><extra><text /></extra><display><text>HIDE</text></display></testcase></testcases></question><question type="coderunner"><name><text>Q5 — Parcours en largeur</text></name><questiontext format="html"><text>&lt;div style="background:#e4f1f2;border:1px solid #b7dadd;border-left:5px solid #026573;border-radius:8px;padding:20px 24px;margin-bottom:18px;"&gt;&lt;p style="color:#026573;font-family:Georgia,serif;font-size:13px;font-weight:bold;letter-spacing:2px;text-transform:uppercase;margin:0 0 10px 0;"&gt;Q5 &amp;mdash; Parcours en largeur&lt;/p&gt;&lt;p style="color:#1e293b;font-family:Georgia,serif;font-size:15px;line-height:1.8;margin:0;"&gt;Le parcours en largeur visite les nœuds niveau par niveau à l'aide d'une file (FIFO).&lt;/p&gt;&lt;/div&gt;&lt;p style="font-family:Georgia,serif;font-size:15px;color:#1e293b;line-height:1.8;"&gt;Écrire &lt;strong&gt;&lt;code&gt;largeur(a)&lt;/code&gt;&lt;/strong&gt; qui renvoie la liste des valeurs dans l'ordre du parcours en largeur. On pourra utiliser &lt;code&gt;collections.deque&lt;/code&gt;.&lt;/p&gt;&lt;div style="background:#f4f9fc;border:1px solid #b7dadd;border-left:5px solid #026573;border-radius:8px;padding:14px 20px;margin:14px 0;"&gt;&lt;pre style="color:#0a3a40;font-family:'Courier New',monospace;font-size:14px;margin:0;background:none;border:none;"&gt;largeur([4,[2,None,None],[6,None,None]]) &amp;rarr; [4, 2, 6]&lt;/pre&gt;&lt;/div&gt;&lt;p style="font-family:Georgia,serif;font-size:14px;color:#475569;line-height:1.7;"&gt;Convention : un arbre vaut &lt;code&gt;None&lt;/code&gt; (vide) ou &lt;code&gt;[valeur, gauche, droite]&lt;/code&gt;. Ex. &lt;code&gt;[4, [2, None, None], [6, None, None]]&lt;/code&gt;.&lt;/p&gt;</text></questiontext><generalfeedback format="html"><text /></generalfeedback><defaultgrade>2.0000000</defaultgrade><penalty>0.0000000</penalty><hidden>0</hidden><idnumber /><coderunnertype>python3</coderunnertype><prototypetype>0</prototypetype><allornothing>0</allornothing><penaltyregime>0</penaltyregime><precheck>0</precheck><hidecheck>0</hidecheck><showsource>0</showsource><answerboxlines>18</answerboxlines><answerboxcolumns>100</answerboxcolumns><answerpreload>from collections import deque
def largeur(a):
    # enfiler la racine, defiler, enfiler les fils
    ...
</answerpreload><globalextra /><useace /><resultcolumns /><template /><iscombinatortemplate /><allowmultiplestdins /><answer>from collections import deque
def largeur(a):
    if a is None:
        return []
    file = deque([a])
    ordre = []
    while file:
        n = file.popleft()
        ordre.append(n[0])
        if n[1] is not None:
            file.append(n[1])
        if n[2] is not None:
            file.append(n[2])
    return ordre
</answer><validateonsave>1</validateonsave><testsplitterre /><language /><acelang /><sandbox /><grader /><cputimelimitsecs /><memlimitmb /><sandboxparams /><templateparams /><hoisttemplateparams>1</hoisttemplateparams><extractcodefromjson>1</extractcodefromjson><templateparamslang>None</templateparamslang><templateparamsevalpertry>0</templateparamsevalpertry><templateparamsevald>{}</templateparamsevald><twigall>0</twigall><uiplugin /><uiparameters /><attachments>0</attachments><attachmentsrequired>0</attachmentsrequired><maxfilesize>10240</maxfilesize><filenamesregex /><filenamesexplain /><displayfeedback>1</displayfeedback><giveupallowed>0</giveupallowed><prototypeextra /><testcases><testcase testtype="0" useasexample="1" hiderestiffail="0" mark="1.0000000"><testcode><text>a = [4, [2, [1,None,None], [3,None,None]], [6, [5,None,None], [7,None,None]]]
print(largeur(a))
</text></testcode><stdin><text /></stdin><expected><text>[4, 2, 6, 1, 3, 5, 7]
</text></expected><extra><text /></extra><display><text>SHOW</text></display></testcase><testcase testtype="0" useasexample="1" hiderestiffail="0" mark="1.0000000"><testcode><text>print(largeur([5,None,None]))
</text></testcode><stdin><text /></stdin><expected><text>[5]
</text></expected><extra><text /></extra><display><text>SHOW</text></display></testcase><testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000"><testcode><text>print(largeur(None))
</text></testcode><stdin><text /></stdin><expected><text>[]
</text></expected><extra><text /></extra><display><text>SHOW</text></display></testcase><testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000"><testcode><text>print(largeur([1,[2,None,None],[3,[4,None,None],None]]))
</text></testcode><stdin><text /></stdin><expected><text>[1, 2, 3, 4]
</text></expected><extra><text /></extra><display><text>HIDE</text></display></testcase></testcases></question></quiz>