<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name><text>IA Q1 &#8212; Distance euclidienne</text></name>
    <questiontext format="html"><text><![CDATA[<p>En KNN, la proximité de deux points se mesure par la distance euclidienne \(d(A,B)=\sqrt{(x_B-x_A)^2+(y_B-y_A)^2}\).</p><p>On considère \(A=(0,0)\) et \(B=(1,2)\). Donner la valeur <strong>exacte</strong> de \(d(A,B)\) (forme avec racine, sans valeur décimale).</p><p>\(d(A,B) = \) [[input:ans1]] [[validation:ans1]]</p>]]></text></questiontext>
    <generalfeedback format="html"><text><![CDATA[<p>\(d(A,B)=\sqrt{(1-0)^2+(2-0)^2}=\sqrt{1+4}=\sqrt{5}\approx 2{,}236\).</p>]]></text></generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.1000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion><text>2023100900</text></stackversion>
    <questionvariables><text>ta : sqrt(5);</text></questionvariables>
    <specificfeedback format="html"><text><![CDATA[<p>[[feedback:prt1]]</p>]]></text></specificfeedback>
    <questionnote><text>d=sqrt(5)</text></questionnote>
    <questiondescription format="html"><text></text></questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html"><text><![CDATA[<p>Correct, bravo&#8239;!</p>]]></text></prtcorrect>
    <prtpartiallycorrect format="html"><text><![CDATA[<p>Partiellement correct.</p>]]></text></prtpartiallycorrect>
    <prtincorrect format="html"><text><![CDATA[<p>Incorrect.</p>]]></text></prtincorrect>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta</tans>
      <boxsize>12</boxsize>
      <strictsyntax>0</strictsyntax>
      <insertstars>1</insertstars>
      <syntaxhint>sqrt(5)</syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options></options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables><text></text></feedbackvariables>
      <node>
        <name>0</name>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html"><text><![CDATA[<p>Oui&#8239;: \(\sqrt{1+4}=\sqrt5\).</p>]]></text></truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html"><text><![CDATA[<p>Rappel&#8239;: \(d=\sqrt{(x_B-x_A)^2+(y_B-y_A)^2}\).</p>]]></text></falsefeedback>
      </node>
    </prt>
  </question>
  <question type="stack">
    <name><text>IA Q2 &#8212; Exactitude (matrice de confusion)</text></name>
    <questiontext format="html"><text><![CDATA[<p>Un classifieur a la matrice de confusion : \(VP=40\), \(VN=30\), \(FP=10\), \(FN=20\).</p><p>Donner l'<strong>exactitude</strong> \(\dfrac{VP+VN}{VP+VN+FP+FN}\) sous forme de fraction irréductible.</p><p>accuracy \(=\) [[input:ans1]] [[validation:ans1]]</p>]]></text></questiontext>
    <generalfeedback format="html"><text><![CDATA[<p>\(\text{accuracy}=\dfrac{40+30}{100}=\dfrac{70}{100}=\dfrac{7}{10}\).</p>]]></text></generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.1000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion><text>2023100900</text></stackversion>
    <questionvariables><text>ta : 7/10;</text></questionvariables>
    <specificfeedback format="html"><text><![CDATA[<p>[[feedback:prt1]]</p>]]></text></specificfeedback>
    <questionnote><text>accuracy=7/10</text></questionnote>
    <questiondescription format="html"><text></text></questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html"><text><![CDATA[<p>Correct, bravo&#8239;!</p>]]></text></prtcorrect>
    <prtpartiallycorrect format="html"><text><![CDATA[<p>Partiellement correct.</p>]]></text></prtpartiallycorrect>
    <prtincorrect format="html"><text><![CDATA[<p>Incorrect.</p>]]></text></prtincorrect>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta</tans>
      <boxsize>10</boxsize>
      <strictsyntax>0</strictsyntax>
      <insertstars>1</insertstars>
      <syntaxhint>7/10</syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options></options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables><text></text></feedbackvariables>
      <node>
        <name>0</name>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html"><text><![CDATA[<p>Oui&#8239;: \(\dfrac{70}{100}=\dfrac{7}{10}\).</p>]]></text></truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html"><text><![CDATA[<p>accuracy \(=\dfrac{VP+VN}{\text{total}}\).</p>]]></text></falsefeedback>
      </node>
    </prt>
  </question>
  <question type="stack">
    <name><text>IA Q3 &#8212; Pr&#233;cision</text></name>
    <questiontext format="html"><text><![CDATA[<p>Avec \(VP=40\) et \(FP=10\), donner la <strong>précision</strong> \(\dfrac{VP}{VP+FP}\) (fraction irréductible).</p><p>précision \(=\) [[input:ans1]] [[validation:ans1]]</p>]]></text></questiontext>
    <generalfeedback format="html"><text><![CDATA[<p>\(\text{précision}=\dfrac{40}{40+10}=\dfrac{40}{50}=\dfrac{4}{5}\).</p>]]></text></generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.1000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion><text>2023100900</text></stackversion>
    <questionvariables><text>ta : 4/5;</text></questionvariables>
    <specificfeedback format="html"><text><![CDATA[<p>[[feedback:prt1]]</p>]]></text></specificfeedback>
    <questionnote><text>precision=4/5</text></questionnote>
    <questiondescription format="html"><text></text></questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html"><text><![CDATA[<p>Correct, bravo&#8239;!</p>]]></text></prtcorrect>
    <prtpartiallycorrect format="html"><text><![CDATA[<p>Partiellement correct.</p>]]></text></prtpartiallycorrect>
    <prtincorrect format="html"><text><![CDATA[<p>Incorrect.</p>]]></text></prtincorrect>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta</tans>
      <boxsize>10</boxsize>
      <strictsyntax>0</strictsyntax>
      <insertstars>1</insertstars>
      <syntaxhint>4/5</syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options></options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables><text></text></feedbackvariables>
      <node>
        <name>0</name>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html"><text><![CDATA[<p>Oui&#8239;: \(\dfrac{40}{50}=\dfrac{4}{5}\).</p>]]></text></truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html"><text><![CDATA[<p>précision \(=\dfrac{VP}{VP+FP}\).</p>]]></text></falsefeedback>
      </node>
    </prt>
  </question>
  <question type="stack">
    <name><text>IA Q4 &#8212; Rappel</text></name>
    <questiontext format="html"><text><![CDATA[<p>Avec \(VP=40\) et \(FN=20\), donner le <strong>rappel</strong> \(\dfrac{VP}{VP+FN}\) (fraction irréductible).</p><p>rappel \(=\) [[input:ans1]] [[validation:ans1]]</p>]]></text></questiontext>
    <generalfeedback format="html"><text><![CDATA[<p>\(\text{rappel}=\dfrac{40}{40+20}=\dfrac{40}{60}=\dfrac{2}{3}\).</p>]]></text></generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.1000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion><text>2023100900</text></stackversion>
    <questionvariables><text>ta : 2/3;</text></questionvariables>
    <specificfeedback format="html"><text><![CDATA[<p>[[feedback:prt1]]</p>]]></text></specificfeedback>
    <questionnote><text>rappel=2/3</text></questionnote>
    <questiondescription format="html"><text></text></questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html"><text><![CDATA[<p>Correct, bravo&#8239;!</p>]]></text></prtcorrect>
    <prtpartiallycorrect format="html"><text><![CDATA[<p>Partiellement correct.</p>]]></text></prtpartiallycorrect>
    <prtincorrect format="html"><text><![CDATA[<p>Incorrect.</p>]]></text></prtincorrect>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta</tans>
      <boxsize>10</boxsize>
      <strictsyntax>0</strictsyntax>
      <insertstars>1</insertstars>
      <syntaxhint>2/3</syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options></options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables><text></text></feedbackvariables>
      <node>
        <name>0</name>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html"><text><![CDATA[<p>Oui&#8239;: \(\dfrac{40}{60}=\dfrac{2}{3}\).</p>]]></text></truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html"><text><![CDATA[<p>rappel \(=\dfrac{VP}{VP+FN}\).</p>]]></text></falsefeedback>
      </node>
    </prt>
  </question>
</quiz>
