<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name><text>Graphes Q1 &#8212; Nombre maximal d&#x27;ar&#234;tes (non orient&#233;)</text></name>
    <questiontext format="html"><text><![CDATA[<p>On considère un <strong>graphe simple non orienté</strong> d'ordre \(n = 6\) (ni boucle, ni multi-arête).</p><p>Quel est le <strong>nombre maximal d'arêtes</strong> de ce graphe&#8239;?</p><p>Nombre maximal d'arêtes \(=\) [[input:ans1]] [[validation:ans1]]</p>]]></text></questiontext>
    <generalfeedback format="html"><text><![CDATA[<p>Chaque paire de sommets distincts porte au plus une arête&#8239;: il y a \(\binom{n}{2} = \dfrac{n(n-1)}{2}\) paires. Pour \(n = 6\)&#8239;: \(\dfrac{6\times 5}{2} = 15\).</p>]]></text></generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.1000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion><text>2023100900</text></stackversion>
    <questionvariables><text>ta : 6*5/2;</text></questionvariables>
    <specificfeedback format="html"><text><![CDATA[<p>[[feedback:prt1]]</p>]]></text></specificfeedback>
    <questionnote><text>max aretes=15</text></questionnote>
    <questiondescription format="html"><text></text></questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html"><text><![CDATA[<p>Correct, bravo&#8239;!</p>]]></text></prtcorrect>
    <prtpartiallycorrect format="html"><text><![CDATA[<p>Partiellement correct.</p>]]></text></prtpartiallycorrect>
    <prtincorrect format="html"><text><![CDATA[<p>Incorrect.</p>]]></text></prtincorrect>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta</tans>
      <boxsize>10</boxsize>
      <strictsyntax>0</strictsyntax>
      <insertstars>1</insertstars>
      <syntaxhint>15</syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options></options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables><text></text></feedbackvariables>
      <node>
        <name>0</name>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html"><text><![CDATA[<p>Oui&#8239;: \(\dfrac{6\times 5}{2} = 15\).</p>]]></text></truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html"><text><![CDATA[<p>Rappel&#8239;: le maximum est \(\dfrac{n(n-1)}{2}\) (nombre de paires de sommets).</p>]]></text></falsefeedback>
      </node>
    </prt>
  </question>
  <question type="stack">
    <name><text>Graphes Q2 &#8212; Nombre maximal d&#x27;arcs (orient&#233;)</text></name>
    <questiontext format="html"><text><![CDATA[<p>On considère un <strong>graphe simple orienté</strong> d'ordre \(n = 5\) (sans boucle).</p><p>Quel est le <strong>nombre maximal d'arcs</strong>&#8239;?</p><p>Nombre maximal d'arcs \(=\) [[input:ans1]] [[validation:ans1]]</p>]]></text></questiontext>
    <generalfeedback format="html"><text><![CDATA[<p>Pour chaque couple ordonné \((u,v)\) de sommets distincts on peut placer un arc \(u\to v\)&#8239;: il y a \(n(n-1)\) couples. Pour \(n = 5\)&#8239;: \(5\times 4 = 20\).</p>]]></text></generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.1000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion><text>2023100900</text></stackversion>
    <questionvariables><text>ta : 5*4;</text></questionvariables>
    <specificfeedback format="html"><text><![CDATA[<p>[[feedback:prt1]]</p>]]></text></specificfeedback>
    <questionnote><text>max arcs=20</text></questionnote>
    <questiondescription format="html"><text></text></questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html"><text><![CDATA[<p>Correct, bravo&#8239;!</p>]]></text></prtcorrect>
    <prtpartiallycorrect format="html"><text><![CDATA[<p>Partiellement correct.</p>]]></text></prtpartiallycorrect>
    <prtincorrect format="html"><text><![CDATA[<p>Incorrect.</p>]]></text></prtincorrect>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta</tans>
      <boxsize>10</boxsize>
      <strictsyntax>0</strictsyntax>
      <insertstars>1</insertstars>
      <syntaxhint>20</syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options></options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables><text></text></feedbackvariables>
      <node>
        <name>0</name>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html"><text><![CDATA[<p>Oui&#8239;: \(n(n-1) = 5\times 4 = 20\).</p>]]></text></truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html"><text><![CDATA[<p>Pour un graphe orienté simple le maximum est \(n(n-1)\), le double du cas non orienté.</p>]]></text></falsefeedback>
      </node>
    </prt>
  </question>
  <question type="stack">
    <name><text>Graphes Q3 &#8212; Lemme des poign&#233;es de main</text></name>
    <questiontext format="html"><text><![CDATA[<p>Un graphe non orienté possède \(m = 9\) arêtes.</p><p>D'après le lemme des poignées de main, que vaut la <strong>somme des degrés</strong> de tous les sommets&#8239;?</p><p>\(\sum_u \deg(u) =\) [[input:ans1]] [[validation:ans1]]</p>]]></text></questiontext>
    <generalfeedback format="html"><text><![CDATA[<p>Chaque arête est incidente à deux sommets&#8239;: elle compte \(1\) dans le degré de chacune de ses extrémités, soit \(2\) au total. Donc \(\sum_u \deg(u) = 2m = 2\times 9 = 18\).</p>]]></text></generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.1000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion><text>2023100900</text></stackversion>
    <questionvariables><text>ta : 2*9;</text></questionvariables>
    <specificfeedback format="html"><text><![CDATA[<p>[[feedback:prt1]]</p>]]></text></specificfeedback>
    <questionnote><text>somme degres=18</text></questionnote>
    <questiondescription format="html"><text></text></questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html"><text><![CDATA[<p>Correct, bravo&#8239;!</p>]]></text></prtcorrect>
    <prtpartiallycorrect format="html"><text><![CDATA[<p>Partiellement correct.</p>]]></text></prtpartiallycorrect>
    <prtincorrect format="html"><text><![CDATA[<p>Incorrect.</p>]]></text></prtincorrect>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta</tans>
      <boxsize>10</boxsize>
      <strictsyntax>0</strictsyntax>
      <insertstars>1</insertstars>
      <syntaxhint>18</syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options></options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables><text></text></feedbackvariables>
      <node>
        <name>0</name>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html"><text><![CDATA[<p>Oui&#8239;: \(2m = 18\).</p>]]></text></truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html"><text><![CDATA[<p>La somme des degrés vaut \(2m\) (chaque arête compte deux fois).</p>]]></text></falsefeedback>
      </node>
    </prt>
  </question>
  <question type="stack">
    <name><text>Graphes Q4 &#8212; Nombre d&#x27;ar&#234;tes &#224; partir des degr&#233;s</text></name>
    <questiontext format="html"><text><![CDATA[<p>Dans un graphe non orienté, la somme des degrés de tous les sommets vaut \(24\).</p><p>Combien ce graphe possède-t-il d'<strong>arêtes</strong>&#8239;?</p><p>Nombre d'arêtes \(m =\) [[input:ans1]] [[validation:ans1]]</p>]]></text></questiontext>
    <generalfeedback format="html"><text><![CDATA[<p>Le lemme des poignées de main donne \(\sum_u \deg(u) = 2m\), donc \(m = \dfrac{1}{2}\sum_u \deg(u) = \dfrac{24}{2} = 12\).</p>]]></text></generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.1000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion><text>2023100900</text></stackversion>
    <questionvariables><text>ta : 24/2;</text></questionvariables>
    <specificfeedback format="html"><text><![CDATA[<p>[[feedback:prt1]]</p>]]></text></specificfeedback>
    <questionnote><text>m=12</text></questionnote>
    <questiondescription format="html"><text></text></questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html"><text><![CDATA[<p>Correct, bravo&#8239;!</p>]]></text></prtcorrect>
    <prtpartiallycorrect format="html"><text><![CDATA[<p>Partiellement correct.</p>]]></text></prtpartiallycorrect>
    <prtincorrect format="html"><text><![CDATA[<p>Incorrect.</p>]]></text></prtincorrect>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta</tans>
      <boxsize>10</boxsize>
      <strictsyntax>0</strictsyntax>
      <insertstars>1</insertstars>
      <syntaxhint>12</syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options></options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables><text></text></feedbackvariables>
      <node>
        <name>0</name>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html"><text><![CDATA[<p>Oui&#8239;: \(m = \dfrac{24}{2} = 12\).</p>]]></text></truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html"><text><![CDATA[<p>On a \(m = \dfrac{1}{2}\sum_u \deg(u)\).</p>]]></text></falsefeedback>
      </node>
    </prt>
  </question>
  <question type="stack">
    <name><text>Graphes Q5 &#8212; Plus court chemin (Dijkstra)</text></name>
    <questiontext format="html"><text><![CDATA[<p>Soit le graphe orienté pondéré (poids positifs) d'arcs&#8239;: \(A\to B\) (4), \(A\to C\) (1), \(C\to B\) (2), \(B\to D\) (1), \(C\to D\) (5).</p><p>En appliquant l'algorithme de Dijkstra depuis \(A\), quelle est la <strong>distance du plus court chemin de \(A\) à \(D\)</strong>&#8239;?</p><p>\(d(A,D) =\) [[input:ans1]] [[validation:ans1]]</p>]]></text></questiontext>
    <generalfeedback format="html"><text><![CDATA[<p>Dijkstra fige \(d(A)=0\), \(d(C)=1\), puis \(d(B)=\min(4,\,1+2)=3\), enfin \(d(D)=\min(1+5,\,3+1)=4\). Le plus court chemin est \(A\to C\to B\to D\) de coût \(1+2+1=4\).</p>]]></text></generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.1000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion><text>2023100900</text></stackversion>
    <questionvariables><text>ta : 4;</text></questionvariables>
    <specificfeedback format="html"><text><![CDATA[<p>[[feedback:prt1]]</p>]]></text></specificfeedback>
    <questionnote><text>d(A,D)=4</text></questionnote>
    <questiondescription format="html"><text></text></questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html"><text><![CDATA[<p>Correct, bravo&#8239;!</p>]]></text></prtcorrect>
    <prtpartiallycorrect format="html"><text><![CDATA[<p>Partiellement correct.</p>]]></text></prtpartiallycorrect>
    <prtincorrect format="html"><text><![CDATA[<p>Incorrect.</p>]]></text></prtincorrect>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta</tans>
      <boxsize>10</boxsize>
      <strictsyntax>0</strictsyntax>
      <insertstars>1</insertstars>
      <syntaxhint>4</syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options></options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables><text></text></feedbackvariables>
      <node>
        <name>0</name>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html"><text><![CDATA[<p>Oui&#8239;: le chemin \(A\to C\to B\to D\) coûte 4.</p>]]></text></truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html"><text><![CDATA[<p>Attention&#8239;: \(A\to B\to D\) coûte \(4+1=5\)&#8239;; passer par \(C\) est meilleur.</p>]]></text></falsefeedback>
      </node>
    </prt>
  </question>
</quiz>
